# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(L):

  """
  The location of the longest string is returned by this function given a list of strings.
  >>> longest(['happy', 'birthday', 'to', 'me'])
  1
  >>> longest(['abc','def','ghi'])
  0
  """

  long_str = ""
  pos = -1

  for i in range(len(L)):
    if len(L[i]) > len(long_str):
      long_str = L[i]
      pos = i

  return pos


### END SOLUTION
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest(['abc','def','ghi'])
Expecting:
    0
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION

def f_to_c(tempf):
  """
  Transforms a temperature's degree Celsius to degree Fahrenheit.
  Args: tempf (float): Fahrenheit temperature in degrees.
  float: Temperature in Celsius is returned.
  Examples:
  >>> f_to_c(212)
  100.0
  >>> f_to_c(32)
  0.0
  >>> f_to_c(98.6)
  37.0
  """
  tempc = 5/9 * (tempf - 32)
  return tempc

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION

def los(admit_date, discharge_date, primary_diagnosis):
  """
  Determine the number of days, in full days, between the date of admission and the date of discharge.
  The period of stay is always returned as zero when "Observation" is the diagnosis.
  Args:
  admit_date (str): The admission date in the 'YYYY-MM-DD' format.
  discharge_date (str): The discharge date in the 'YYYY-MM-DD' format.
  Primary diagnosis is the diagnosis.
  Returns: int: Total number of days spent there.
  Example:
  >>> los(date(2019,1,3), date(2019,1,8), 'COPD')
  5
  >>> los(date(2020,1,5), date(2020,1,8), 'Observation')
  0
  """
  if primary_diagnosis == "Observation":
    return 0
  length_of_stay_days = (discharge_date - admit_date).days
  return length_of_stay_days

### END SOLUTION
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), 'COPD')
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), 'Observation')
Expecting:
    0
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python

### BEGIN SOLUTION

def average_los(encounters):
  """
  Determine the typical Length of Stay (LOS) for interactions.
  encounters (dict): Dictionary of encounters, each of which is represented by a dictionary having at least two dates—the admit date and the discharge date—in it.
  float: Average Length of Stay (LOS) is what is returned.
  Example:

  encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
  }
  >>> average_los(encounters)
  3.3333333333333335
  """
  total_length_of_stay = 0
  length_of_pat = 0
  for encid, visit in encounters.items():
    length_of_stay = los(visit['admit'],visit['discharge'],visit['diagnosis'])
    if (length_of_stay > 0):
      total_length_of_stay += length_of_stay
      length_of_pat += 1
  return total_length_of_stay / length_of_pat

### END SOLUTION
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    average_los(encounters)
Expecting:
    3.3333333333333335
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f(cel):
  """
  Fahrenheit to degrees Celsius conversion.
  Temperature in Celsius is indicated by the argument celsius (float).
  Returns:
  float: The current temperature in Fahrenheit.
  Examples:
  >>> c_to_f(100)
  212.0
  >>> c_to_f(0)
  32.0
  >>> c_to_f(37)
  98.6
  """
  temp_f = (cel * 9/5) + 32
  return temp_f

### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
```

```python
!git config --global user.email 'tannerunaveen4@gmail.com'
!git config --global user.name 'tannerunaveen4'
```

```python
!git clone https://github.com/tannerunaveen4/hds5210-2023.git
```

Output:
```
fatal: destination path 'hds5210-2023' already exists and is not an empty directory.
```

```python
%cd /content/hds5210-2023
```

Output:
```
/content/hds5210-2023
```

```python
%ls
```

Output:
```
Copy_of_week02_assignment_1.ipynb  README.md  [0m[01;34mweek02[0m/  [01;34mweek04[0m/
[01;34mhds5210-2023[0m/                      [01;34mweek01[0m/    [01;34mweek03[0m/  [01;34mweek05[0m/
```

```python
%cd /content/hds5210-2023/week05
```

Output:
```
/content/hds5210-2023/week05
```

```python
%ls
```

Output:
```
week05_assignment_1.ipynb  week05_assignment_2.ipynb
```

```python
%cd ..
```

Output:
```
/content/hds5210-2023
```

```python
%ls
```

Output:
```
Copy_of_week02_assignment_1.ipynb  README.md  [0m[01;34mweek02[0m/  [01;34mweek04[0m/
[01;34mhds5210-2023[0m/                      [01;34mweek01[0m/    [01;34mweek03[0m/  [01;34mweek05[0m/
```

```python
!git status
```

Output:
```
On branch main
Your branch is ahead of 'origin/main' by 2 commits.
  (use "git push" to publish your local commits)

nothing to commit, working tree clean
```

```python
!git add --all
```

```python
!git commit -a -m "Week 5 assignment two"
```

Output:
```
On branch main
Your branch is ahead of 'origin/main' by 2 commits.
  (use "git push" to publish your local commits)

nothing to commit, working tree clean
```

```python
!git remote -v
```

Output:
```
origin	https://github.com/tannerunaveen4/hds5210-2023.git (fetch)
origin	https://github.com/tannerunaveen4/hds5210-2023.git (push)
```

```python
!git push https://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@github.com/tannerunaveen4/hds5210-2023.git
```

Output:
```
Everything up-to-date
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.