<a href="https://colab.research.google.com/github/tannerunaveen4/hds5210-2023/blob/main/week14/week14_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.  All of these files are found in our usual s3 bucket: `https://hds5210-data.s3.amazonaws.com`

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.

* https://hds5210-data.s3.amazonaws.com/npidata.csv


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

* https://hds5210-data.s3.amazonaws.com/cms2014.csv
* https://hds5210-data.s3.amazonaws.com/cms2015.csv
* https://hds5210-data.s3.amazonaws.com/cms2016.csv

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

```python
cms2014 = pd.read_csv("https://hds5210-data.s3.amazonaws.com/cms2014.csv")
#cms2014.shape #(63093, 59)
cms2014.columns
```

Output:
```
Index(['nbr', 'npi', 'provider_type', 'medicare_participation_indicator',
       'number_of_hcpcs', 'total_services', 'total_unique_benes',
       'total_submitted_chrg_amt', 'total_medicare_allowed_amt',
       'total_medicare_payment_amt', 'total_medicare_stnd_amt',
       'drug_suppress_indicator', 'number_of_drug_hcpcs',
       'total_drug_services', 'total_drug_unique_benes',
       'total_drug_submitted_chrg_amt', 'total_drug_medicare_allowed_amt',
       'total_drug_medicare_payment_amt', 'total_drug_medicare_stnd_amt',
       'med_suppress_indicator', 'number_of_med_hcpcs', 'total_med_services',
       'total_med_unique_benes', 'total_med_submitted_chrg_amt',
       'total_med_medicare_allowed_amt', 'total_med_medicare_payment_amt',
       'total_med_medicare_stnd_amt', 'beneficiary_average_age',
       'beneficiary_age_less_65_count', 'beneficiary_age_65_74_count',
       'beneficiary_age_75_84_count', 'beneficiary_age_greater_84_count',
       'beneficiary_female_count', 'ben
```

```python
cms2015 = pd.read_csv("https://hds5210-data.s3.amazonaws.com/cms2015.csv")
#cms2015.shape #(64540, 59)
cms2015.columns
```

Output:
```
Index(['nbr', 'npi', 'provider_type', 'medicare_participation_indicator',
       'number_of_hcpcs', 'total_services', 'total_unique_benes',
       'total_submitted_chrg_amt', 'total_medicare_allowed_amt',
       'total_medicare_payment_amt', 'total_medicare_stnd_amt',
       'drug_suppress_indicator', 'number_of_drug_hcpcs',
       'total_drug_services', 'total_drug_unique_benes',
       'total_drug_submitted_chrg_amt', 'total_drug_medicare_allowed_amt',
       'total_drug_medicare_payment_amt', 'total_drug_medicare_stnd_amt',
       'med_suppress_indicator', 'number_of_med_hcpcs', 'total_med_services',
       'total_med_unique_benes', 'total_med_submitted_chrg_amt',
       'total_med_medicare_allowed_amt', 'total_med_medicare_payment_amt',
       'total_med_medicare_stnd_amt', 'beneficiary_average_age',
       'beneficiary_age_less_65_count', 'beneficiary_age_65_74_count',
       'beneficiary_age_75_84_count', 'beneficiary_age_greater_84_count',
       'beneficiary_female_count', 'ben
```

```python
cms2016 = pd.read_csv("https://hds5210-data.s3.amazonaws.com/cms2016.csv")
#cms2016.shape #(66229, 59)
cms2016.columns
```

Output:
```
Index(['nbr', 'NPI', 'PROVIDER_TYPE', 'MEDICARE_PARTICIPATION_INDICATOR',
       'NUMBER_OF_HCPCS', 'TOTAL_SERVICES', 'TOTAL_UNIQUE_BENES',
       'TOTAL_SUBMITTED_CHRG_AMT', 'TOTAL_MEDICARE_ALLOWED_AMT',
       'TOTAL_MEDICARE_PAYMENT_AMT', 'TOTAL_MEDICARE_STND_AMT',
       'DRUG_SUPPRESS_INDICATOR', 'NUMBER_OF_DRUG_HCPCS',
       'TOTAL_DRUG_SERVICES', 'TOTAL_DRUG_UNIQUE_BENES',
       'TOTAL_DRUG_SUBMITTED_CHRG_AMT', 'TOTAL_DRUG_MEDICARE_ALLOWED_AMT',
       'TOTAL_DRUG_MEDICARE_PAYMENT_AMT', 'TOTAL_DRUG_MEDICARE_STND_AMT',
       'MED_SUPPRESS_INDICATOR', 'NUMBER_OF_MED_HCPCS', 'TOTAL_MED_SERVICES',
       'TOTAL_MED_UNIQUE_BENES', 'TOTAL_MED_SUBMITTED_CHRG_AMT',
       'TOTAL_MED_MEDICARE_ALLOWED_AMT', 'TOTAL_MED_MEDICARE_PAYMENT_AMT',
       'TOTAL_MED_MEDICARE_STND_AMT', 'BENEFICIARY_AVERAGE_AGE',
       'BENEFICIARY_AGE_LESS_65_COUNT', 'BENEFICIARY_AGE_65_74_COUNT',
       'BENEFICIARY_AGE_75_84_COUNT', 'BENEFICIARY_AGE_GREATER_84_COUNT',
       'BENEFICIARY_FEMALE_COUNT', 'BEN
```

```python
cms2016.columns = cms2016.columns.str.lower() if not all(cms2016.columns == cms2016.columns.str.lower()) else cms2016.columns
cms2014.columns = cms2014.columns.str.lower() if not all(cms2014.columns == cms2016.columns.str.lower()) else cms2014.columns
cms2015.columns = cms2015.columns.str.lower() if not all(cms2015.columns == cms2016.columns.str.lower()) else cms2015.columns
cms2016.columns
```

Output:
```
Index(['nbr', 'npi', 'provider_type', 'medicare_participation_indicator',
       'number_of_hcpcs', 'total_services', 'total_unique_benes',
       'total_submitted_chrg_amt', 'total_medicare_allowed_amt',
       'total_medicare_payment_amt', 'total_medicare_stnd_amt',
       'drug_suppress_indicator', 'number_of_drug_hcpcs',
       'total_drug_services', 'total_drug_unique_benes',
       'total_drug_submitted_chrg_amt', 'total_drug_medicare_allowed_amt',
       'total_drug_medicare_payment_amt', 'total_drug_medicare_stnd_amt',
       'med_suppress_indicator', 'number_of_med_hcpcs', 'total_med_services',
       'total_med_unique_benes', 'total_med_submitted_chrg_amt',
       'total_med_medicare_allowed_amt', 'total_med_medicare_payment_amt',
       'total_med_medicare_stnd_amt', 'beneficiary_average_age',
       'beneficiary_age_less_65_count', 'beneficiary_age_65_74_count',
       'beneficiary_age_75_84_count', 'beneficiary_age_greater_84_count',
       'beneficiary_female_count', 'ben
```

```python
cms2016.shape
```

Output:
```
(66229, 59)
```

```python
cms2014['year'] = 2014
cms2015['year'] = 2015
cms2016['year'] = 2016
```

```python
cms2014.shape
```

Output:
```
(63093, 60)
```

```python
cms2015.shape
```

Output:
```
(64540, 60)
```

```python
merged = pd.concat([cms2014, cms2015], keys=['2014', '2015'])
merged.shape
```

Output:
```
(127633, 60)
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
answer = pd.concat([cms2014, cms2015, cms2016], keys=['2014', '2015', '2016'])
answer.shape
### SOLUTION
```

Output:
```
(193862, 60)
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **https://hds5210-data.s3.amazonaws.com/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

```python
npdata = pd.read_csv("https://hds5210-data.s3.amazonaws.com/npidata.csv")
npdata.shape
```

Output:
```
<ipython-input-234-d3f280caed41>:1: DtypeWarning: Columns (18,33) have mixed types. Specify dtype option on import or set low_memory=False.
  npdata = pd.read_csv("https://hds5210-data.s3.amazonaws.com/npidata.csv")
(251661, 43)
```

```python
npdata.columns
```

Output:
```
Index(['row', 'NPI', 'Entity Type Code', 'Replacement NPI',
       'Employer Identification Number (EIN)',
       'Provider Organization Name (Legal Business Name)',
       'Provider Last Name (Legal Name)', 'Provider First Name',
       'Provider Middle Name', 'Provider Name Prefix Text',
       'Provider Name Suffix Text', 'Provider Credential Text',
       'Provider Other Organization Name',
       'Provider Other Organization Name Type Code',
       'Provider Other Last Name', 'Provider Other First Name',
       'Provider Other Middle Name', 'Provider Other Name Prefix Text',
       'Provider Other Name Suffix Text', 'Provider Other Credential Text',
       'Provider Other Last Name Type Code',
       'Provider First Line Business Mailing Address',
       'Provider Second Line Business Mailing Address',
       'Provider Business Mailing Address City Name',
       'Provider Business Mailing Address State Name',
       'Provider Business Mailing Address Postal Code',
       'Provider
```

```python
# Check for duplicate columns
print("Duplicate columns in cms:", cms.columns[cms.columns.duplicated()])
print("Duplicate columns in npdata:", npdata.columns[npdata.columns.duplicated()])

# Check data types of the 'NPI' column
print("Data type of 'NPI' in cms:", cms['NPI'].dtype)

```

Output:
```
Duplicate columns in cms: Index([], dtype='object')
Duplicate columns in npdata: Index([], dtype='object')
Data type of 'NPI' in cms: int64
```

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
cms = pd.concat([cms2014, cms2015, cms2016], keys=['2014', '2015', '2016'])
cms.rename(columns={"npi": "NPI"}, inplace=True)
answer = pd.merge(cms, npdata, on='NPI', how='left')
### SOLUTION
answer.shape
```

Output:
```
(193862, 102)
```

```python
answer.columns
```

Output:
```
Index(['nbr', 'NPI', 'provider_type', 'medicare_participation_indicator',
       'number_of_hcpcs', 'total_services', 'total_unique_benes',
       'total_submitted_chrg_amt', 'total_medicare_allowed_amt',
       'total_medicare_payment_amt',
       ...
       'Provider Business Practice Location Address Postal Code',
       'Provider Business Practice Location Address Country Code (If outside U.S.)',
       'Provider Business Practice Location Address Telephone Number',
       'Provider Business Practice Location Address Fax Number',
       'Provider Enumeration Date', 'Last Update Date',
       'NPI Deactivation Reason Code', 'NPI Deactivation Date',
       'NPI Reactivation Date', 'Provider Gender Code'],
      dtype='object', length=102)
```

```python
common_columns = cms.columns.intersection(npdata.columns)
print("Common Columns:", common_columns)
```

Output:
```
Common Columns: Index(['NPI'], dtype='object')
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
#assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION

```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
answer = None

### SOLUTION

```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

```python
%matplotlib inline
answer.plot.bar()
```

```python

```